Enhanced bacterial chemotaxis in confined microchannels: Optimal performance in lane widths matching circular swimming radius

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

The authors show experimentally that, in 2D, bacteria swim up a chemotactic gradient much more effectively when they are in the presence of lateral walls. Systematic experiments identify an optimum for chemotaxis for a channel width of ~8µm, a value close to the average radius of the circle trajectories of the unconfined bacteria in 2D. These chiral circles impose that the bacteria swim preferentially along the right-side wall, which indeed yields chemotaxis in the presence of a chemotactic gradient. These observations are backed by numerical simulations and a geometrical analysis.

Reviewer #3 (Public review):

This paper addresses, through experiment and simulation, the combined effects of bacterial circular swimming near no-slip surfaces and chemotaxis in simple linear gradients. The authors have constructed a microfluidic device in which a gradient of L-aspartate is established, to which bacteria respond while swimming while confined in channels of different widths. There is a clear effect that the chemotactic drift velocity reaches a maximum in channel widths of about 8 microns, similar in size to the circular orbits that would prevail in the absence of side walls. Numerical studies of simplified models confirm this connection.

The experimental aspects of this study are well executed. The design of the microfluidic system is clever in that it allows a kind of "multiplexing" in which all the different channel widths are available to a given sample of bacteria.

The authors have included a useful intuitive explanation of their results via a geometric model of the trajectories. In future work it would be interesting to analyze further the voluminous data on the trajectories of cells by formulating the mathematical problem in terms of a suitable Fokker-Planck equation for the probability distribution of swimming directions. In particular, this might help understand how incipient circular trajectories are interrupted by collisions with the walls and how this relates to enhanced chemotaxis.

The authors argue that these findings may have relevance to a number of physiological and ecological contexts. As these would be characterized by significant heterogeneity in pore sizes and geometries, further work will be necessary to translate the present results to those situations.

Thanks to the referees' input and more work, we think our revised manuscript now meets the high standard of eLife

Recommendations for the authors:

The importance of the circular swimming chirality for the observed phenomenon could be further emphasized by actually using the word "chiral" or "chirality" in the text. Also indicating what would change is swimming were counterclockwise rather then clockwise would help the reader understand the key significance of chirality.

We thank the reviewer for this insightful suggestion. We agree that the chirality of the surface interaction is central to the observed phenomenon and should be explicitly highlighted to improve the reader's understanding.

In response, we have incorporated the terms "chiral" and "chirality" throughout the manuscript (Abstract, Introduction, Results, and Discussion) to emphasize this aspect. Furthermore, we have added a specific explanation in the Results section (the last paragraph of subsection “The cells in the right sidewall region dominated the chemotaxis of E. coli with lane confinements”) detailing the hypothetical scenario of counter-clockwise swimming. We clarify that in such a case, the hydrodynamic interaction would cause cells to veer left, resulting in up-gradient accumulation along the left sidewall rather than the right. We believe these additions significantly improve the clarity of the underlying physical mechanism.

Reviewer #1 (Recommendations for the authors):

I still have several comments that the authors may want to consider for the last version.

- The run and tumble behavior of the cells at the surface remains puzzling and would need some more explanation in the text. Tumbles with no significant reorientation angle amount largely to smooth swimmers. How can a model based on run-and-tumbles be used to explain the difference between LSW and RSW?

We apologize for the lack of clarity regarding the surface run-and-tumble behavior. While it is true that surface tumbles often result in smaller reorientation angles compared to bulk swimming, they are not negligible and play a critical role in the observed asymmetry. As shown in the tumble angle distributions (Fig. 2E and 2F), the probability of a tumble angle exceeding π/2 is approximately 9% for sidewall trajectories and 30% for the middle area. This tumbling behavior leads to differences between the left sidewall (LSW) and right sidewall (RSW) in two key ways:

First, as detailed in our geometric analysis (Fig. 6), running cells following stable clockwise circular paths are geometrically favored to reach the RSW. Because cells moving up-gradient (towards the RSW) experience suppressed tumbling, they maintain these stable circular trajectories and accumulate effectively. Conversely, cells moving down-gradient (towards the LSW) experience enhanced tumbling. These frequent interruptions distort the circular trajectories required to reach the LSW, resulting in fewer bacteria entering the LSW compared to the RSW.

Second, once at the wall, the difference in tumbling frequency dictates retention. Majority of LSW cells are swimming down-gradient (LSW-DG) and thus tumble more frequently, increasing their probability of escaping the wall. Majority of RSW cells are swimming up-gradient (RSW-UG), suppressing tumbles and increasing their residence time at the wall.

The relevant clarifications have been included in the last paragraph of “Results” in the manuscript.

- Figure 5B would need more explanation. I still don't understand the different behaviors for the right and left side walls at small widths. Is it noise really or a more complex behavior? Since most of these calculations are based precisely on the shape of these curves it would be useful to discuss them in more detail.

We apologize for the lack of clarity. The behavior observed at small widths in Figure 5B is not noise; rather, it reflects the idealized nature of our simulation model.

In the simulation, bacteria were modeled as active particles without explicit steric exclusion for the flagella and cell body. Consequently, simulated cells retain the ability to reorient and turn freely even in very narrow lanes (w ≤ 6 μm), allowing the geometric sorting mechanism (which favors the RSW) to function efficiently even at small widths. This is why the simulation shows a distinct difference between LSW and RSW proportions in this regime.

In the experimental reality, however, the finite size of the bacterial body and flagella creates steric hindrance. In narrow channels, this physical constraint restricts the cells' ability to turn, thereby disrupting the circular swimming mechanism required to sort cells into the RSW. As a result, experimental data shows that the proportions of LSW and RSW cells tend to equalize in narrow channels (e.g., w = 6 μm in Fig. 4B), leading to a lower chemotactic drift velocity than predicted by the simulation.

We have added a discussion regarding these steric effects and the deviation at narrow widths to the Results section (the penultimate paragraph of subsection "Simulation of E. coli chemotaxis within lane confinement") in the revised manuscript.

- The importance of the chirality of the circular trajectories, although essential, remains insufficiently mentioned in the text.

We have incorporated the terms "chiral" and "chirality" throughout the manuscript (Abstract, Introduction, Results, and Discussion) to emphasize this aspect. Furthermore, we have added a specific explanation in the Results section (the last paragraph of subsection “The cells in the right sidewall region dominated the chemotaxis of E. coli with lane confinements”) detailing the hypothetical scenario of counter-clockwise swimming.

- It would be useful to color-code the trajectories of Figure 1B and alike with time.

Thank you for the suggestion. Now the trajectories in Fig. 1B have been redrawn. Distinct colors denote individual trajectories, with color intensity darkening to indicate time progression.

Changes from the earlier version to the revised version

Abstract

Earlier version: Understanding bacterial behavior in confined environments is helpful for elucidating microbial ecology and developing strategies to manage bacterial infections. While extensive research has focused on bacterial motility on surfaces and in porous media, chemotaxis in confined spaces remains poorly understood. Here, we investigate the chemotaxis of Escherichia coli within microfluidic lanes under a linear concentration gradient of L-aspartate. We demonstrate that E. coli exhibits significantly enhanced chemotaxis in lanes with sidewalls compared to open surfaces. We attribute this phenomenon primarily to the intrinsic clockwise circular motion of surface-swimming bacteria and the subsequent alignment effect upon collision with the sidewalls. By varying lane widths, we identify that an 8 μm width—approximating the radius of bacterial circular swimming on surfaces—maximizes chemotactic drift velocity. These results are supported by both experimental observations and stochastic simulations, establishing a clear proportional relationship between optimal lane width and the radius of bacterial circular swimming. Further geometric analysis provides an intuitive understanding of this phenomenon. Our results may offer insights into bacterial navigation in complex biological environments such as host tissues and biofilms, providing a preliminary step toward exploring microbial ecology in confined habitats and potential strategies for controlling bacterial infections.

Revised version: Understanding bacterial behavior in confined environments is helpful for elucidating microbial ecology and developing strategies to manage bacterial infections. While extensive research has focused on bacterial motility on surfaces and in porous media, chemotaxis in confined spaces remains poorly understood. Here, we investigate the chemotaxis of Escherichia coli within microfluidic lanes under a linear concentration gradient of L-aspartate. We demonstrate that E. coli exhibits significantly enhanced chemotaxis in lanes with sidewalls compared to open surfaces. We attribute this phenomenon primarily to the intrinsic chiral clockwise circular motion of surface-swimming bacteria and the subsequent alignment effect upon collision with the sidewalls. By varying lane widths, we identify that an 8 μm width—approximating the radius of bacterial circular swimming on surfaces—maximizes chemotactic drift velocity. These results are supported by both experimental observations and stochastic simulations, establishing a clear proportional relationship between optimal lane width and the radius of bacterial circular swimming. Further geometric analysis provides an intuitive understanding of this phenomenon. Our results may offer insights into bacterial navigation in complex biological environments such as host tissues and biofilms, providing a preliminary step toward exploring microbial ecology in confined habitats and potential strategies for controlling bacterial infections.

Introduction

Microorganisms inhabit a variety of complex confined environments. Sperm swimming in the reproductive tract of a female animal (1), bacteria navigating through host cells and tissues (2), cells moving in the wrinkle of biofilm (3), and swarm cells maneuvering in crowded cell clusters (4) are all examples of organisms operating in geometrically confined environments. Furthermore, many diseases have been found to be related to bacterial motion in confinement, such as urinary tract infections and neonatal sepsis.

Earlier version: Previous research has explored the motion of bacteria in confined environments. Studies have shown that flagellated bacteria tend to accumulate near both solid-liquid (5, 6) and air-liquid interfaces (7), and exhibit a specific direction of circular swimming (8). For example, E. coli swim clockwise when observed from above a solid surface (8-11), whereas Caulobacter crescentus move in tight, counter-clockwise circles when viewed from the liquid side. Moreover, Tang et al. found that flagellated bacteria orbit within the thin fluid film around micrometer-sized particles (12). Bacteria also display different motion behaviors in porous media and tubes (13-19), as well as in other types of confinement (20-22). In narrow channels with width less than 4 μm, bacterial growth (23), swimming trajectories (24) and swimming speeds (25) are significantly impacted compared to bulk conditions. Furthermore, studies on bacterial transport in rectangular microchannels under flow have demonstrated that confinement influences both bacterial fluxes and spatial distribution, with cells exhibiting distinct transport dynamics along channel edges (26). However, while motility in confinement is well-documented, the specific rules governing bacterial chemotaxis within geometrically constrained environments remain less understood.

Revised version: Previous research has explored the motion of bacteria in confined environments. Studies have shown that flagellated bacteria tend to accumulate near both solid-liquid (5, 6) and air-liquid interfaces (7), and exhibit a specific chirality in their direction of circular swimming (8). For example, E. coli swim clockwise when observed from above a solid surface (8–11), whereas Caulobacter crescentus move in tight, counter-clockwise circles when viewed from the liquid side. Moreover, Tang et al. found that flagellated bacteria orbit within the thin fluid film around micrometer-sized particles (12). Bacteria also display different motion behaviors in porous media and tubes (13–19), as well as in other types of confinement (20–22). In narrow channels with width less than 4 μm, bacterial growth (23), swimming trajectories (24) and swimming speeds (25) are significantly impacted compared to bulk conditions. Furthermore, studies on bacterial transport in rectangular microchannels under flow have demonstrated that confinement influences both bacterial fluxes and spatial distribution, with cells exhibiting distinct transport dynamics along channel edges (26). However, while motility in confinement is well-documented, the specific rules governing bacterial chemotaxis within geometrically constrained environments remain less understood.

E. coli can sense and move in chemical gradients through a set of chemotactic proteins (27, 28). The transmembrane chemoreceptors detect attractants or repellents and transmit signals into the cell by modulating the autophosphorylation of the histidine kinase CheA. Attractant binding suppresses CheA autophosphorylation, while repellent binding promotes it. This modulation alters the concentration of the phosphorylated response regulator protein, CheY-P (28–31). CheY-P binds to the cytoplasmic domain of flagellar motors, altering the probability of their clockwise rotation (32, 33). Clockwise and counterclockwise rotations of motors lead to reorientation (tumble) and smooth swimming (run) modes of the cell body, respectively (34). The adaptation enzymes CheR and CheB methylate and demethylate the receptors, respectively, mediating sensory adaptation (35, 36).

While the molecular mechanism is well-characterized, the physical environment plays a crucial role in chemotactic efficiency. Bacterial chemotaxis has been observed under various degrees of confinement, often yielding complex behaviors. For instance, in quasi-2D microchannels with a depth of 2 μm, bacteria were found to adjust their run lengths and tumble frequencies to respond to chemical gradients (37). Similarly, studies in 10 μm-depth microchambers revealed that E. coli maintains chemotactic performance in the presence of evenly spaced obstacles (38). While these studies focused on confinement between top and bottom surfaces, recent work by Braham et al showed that bacteria on a single solid surface can exhibit an escape response to high-concentration chemorepellents (250 mM Ni(NO3)2) (39). Conversely, other research has shown that bacterial chemotactic drift velocity decreases sharply near sample chamber surfaces compared to bulk liquid, a phenomenon attributed to circular cell movement at these interfaces (40). This observation suggests that a single surface can significantly impair bacterial chemotactic performance. Yet, a study of chemotaxis in two-dimensional bacterial swarms on semisolid agar surfaces has demonstrated that E. coli can exhibit effective chemotaxis in attractant gradients, especially at high cell densities (4). This apparent contradiction is likely due to cell-cell interactions mitigating surface effects in dense populations. Despite these investigations into motility under surface effects, a comprehensive understanding of how specific confinement geometries modulate chemotaxis remains limited.

Here, we study the chemotaxis of E. coli on surfaces within lanes of different widths, exposed to a linear concentration gradient of L-aspartate using a hydrogel-based microfluidic device. Our findings reveal enhanced chemotactic performance in the presence of sidewalls compared to their absence. Cells in the middle area (MA) exhibited no drift and did not contribute to overall chemotaxis. However, cells near the left sidewall (LSW) and right sidewall (RSW) demonstrated negative and positive drift, respectively, with the positive drift velocity from the RSW exceeding the negative drift velocity from the LSW.

Furthermore, by comparing lanes of different widths within the same stable linear gradient field, we observed that cells exhibited the highest drift velocity in 8 μm-wide lanes, which approximates the peak value in the distribution of circular swimming radii on surfaces. This optimal width for maximal chemotaxis resulted from the largest percentage of cells occupying the right sidewalls. Stochastic simulations of bacterial chemotaxis corroborated that the optimized lane width varies with the bacteria’s circular swimming radius. Our findings have important implications for the screening of swimming organisms and the study of chemotactic behavior in real-world scenarios.

Results

Chemotaxis of E. coli within lane confinements

To study bacterial chemotaxis in confined environments such as tubes or interstitial tissues, we developed a hydrogel-based microfluidic device featuring lanes of different widths. The device design, illustrated in Fig. 1A, incorporates two 2% (w/v) agarose walls inserted between three main channels: a sink channel, a cell channel, and a source channel. These walls prevent cell passage while allowing chemical diffusion. The agarose walls and main channels measure 100 μm and 400 μm in width, respectively, with the device having a depth of 30 μm.

Earlier version: A. Microfluidic device used in the experiment (top view). B. Thirty bacterial trajectories selected from the data of 44-μm-wide lane in gradient assays. These represent a subset of the trajectories analyzed in panel C. C. The relationship between Δx and nΔt calculated from all trajectories in the lanes with a width of 44 μm. The black dots represent the gradient assay (∇c=0.05 μM/μm) with a total of 3206 tracks from 6 movies, while the gray dots represent no gradient (∇c=0) with a total of 4755 tracks from 10 movies. The light-yellow and light-blue shadows represent the standard error of the mean (SEM) of different trajectories in the gradient assays and the control assays, respectively. Linear fitting was performed to obtain Vd = 1.6 ± 0.3 μm/s (black solid line) and Vd = 0.2 ± 0.2 μm/s (gray solid line) for gradient and control assays respectively. Error in drift velocity represents standard deviation.

Revised version: A. Microfluidic device used in the experiment (top view). B. Thirty bacterial trajectories selected from the data of 44-μm-wide lane in gradient assays. Distinct colors denote individual trajectories, with color intensity darkening to indicate time progression. These represent a subset of the trajectories analyzed in panel C. C. The relationship between Δ𝑥 and nΔ𝑡 calculated from all trajectories in the lanes with a width of 44 μm. The black dots represent the gradient assay (𝛻c=0.05 μM/μm) with a total of 3206 tracks from 6 movies, while the gray dots represent no gradient (𝛻𝑐=0) with a total of 4755 tracks from 10 movies. The light-yellow and light-blue shadows represent the standard error of the mean (SEM) of different trajectories in the gradient assays and the control assays, respectively. Linear fitting was performed to obtain 𝑉𝑑 = 1.6 ± 0.3 μm/s (black solid line) and 𝑉𝑑 = 0.2 ± 0.2 μm/s (gray solid line) for gradient and control assays respectively. Error in drift velocity represents standard deviation.

We established a linear concentration gradient (0.05 μM/μm) of L-aspartate, an attractant sensed by the Tar receptor, by flowing motility medium through the sink channel and 50 μM L-aspartate solution (prepared in motility medium) through the source channel at a constant rate of 10 μL/min. The resulting gradient across the cell channel is depicted in Fig. S1.

The lanes of various widths were created using parallel PDMS (polydimethylsiloxane) pillars spaced at different distances. Each lane, approximately 160 μm in length, is highlighted by a yellow rectangular box in Fig. 1A. E. coli cells, confined to the cell channel, sensed the attractant gradient and performed chemotaxis within these lanes of different widths.

To enhance imaging contrast, we introduced a plasmid expressing mCherry fluorescent protein (pTrc99a-mcherry) into wild-type Escherichia coli HCB1, denoting it as HCB1-pTrc99a-mcherry. A high-sensitivity CMOS camera was used to record bacterial movement trajectories on the bottom surfaces of the lanes in the cell channel, where most cells accumulate due to the combined effects of gravity and hydrodynamic wall attraction in the Percoll-free medium.

We recorded the movement of HCB1-pTrc99a-mcherry cells on the liquid-glass surface of a 44-μm-wide lane with and without a L-aspartate concentration gradient, respectively. An example video is shown in Movie S1, and several representative trajectories are illustrated in Fig. 1B.

Earlier version: Considering the potential escape into the third dimension at the edge (26), we calculated the drift velocities vd of the trajectories from both assays along the x-direction (up-gradient) by fitting the relation ⟨x(i + n) − x(i)⟩∼nΔt (n=1, 3, 5, …, 39) with a linear function, the slope of which is the drift velocity. Here, x(i) represents the cell’s x-position in the i-th frame, and Δt denotes the time interval between two frames (0.05 s). This method yields a drift velocity that is independent of potential cell loss into the third dimension. As shown in Fig. 1C, the drift velocity vd was 1.6±0.3 μm/s and 0.2 ±0.2 μm/s (mean±SD) for the gradient and control (no-gradient) assays, respectively. Thus, bacteria in the L-aspartate gradient exhibited obvious chemotaxis on the lane surfaces, contrasting with previous findings on liquid-solid surfaces without sidewalls (40).

Revised version: Considering the potential escape into the third dimension at the edge (26), we calculated the drift velocities 𝑣𝑑 of the trajectories from both assays along the x-direction (up-gradient) by fitting the relation 〈𝑥(𝑖 + 𝑛) − 𝑥(𝑖)〉∼𝑛Δ𝑡 (n=1, 3, 5, …, 39) with a linear function, the slope of which is the drift velocity. Here, 𝑥(𝑖) represents the cell’s 𝑥-position in the i-th frame, and Δ𝑡 denotes the time interval between two frames (0.05 s). This method yields a drift velocity that is independent of potential cell loss into the third dimension. As shown in Fig. 1C, the drift velocity 𝑣𝑑 was 1.6±0.3 μm/s and 0.2 ±0.2 μm/s (mean±SD) for the gradient and control (no-gradient) assays, respectively. Thus, bacteria in the L-aspartate gradient exhibited obvious chemotaxis on the lane surfaces, contrasting with previous findings on liquid-solid surfaces without sidewalls (40).

The cells in the right sidewall region dominated the chemotaxis of E. coli within lane confinements

To elucidate the mechanism of bacterial chemotaxis within lane confinements, we analyzed cell trajectories, identifying three distinct motion states due to the presence of sidewalls. As illustrated in Fig. 2A, cells exhibited straight-line swimming along sidewalls, movement towards or away from sidewalls, or circular swimming. Fig. 2B displays three representative trajectories. Circular swimming is characteristic of typical E. coli surface motion (8). We further analyzed the trajectories of bacteria colliding with the sidewalls. As shown in Fig. S2, cells colliding with the sidewalls tend to align with and swim along the sidewalls, and the alignment timescale is negligible compared to the typical wall residence time (6).

Earlier version: A. Schematic drawing of three typical motion states. B. Trajectories for three typical motion states. Red dots represent the start of each trajectory. C. Calculating the rotational exponent for each trajectory, γR, by fitting the mean-squared orientational displacement ⟨MSOD⟩. Dots were experimental data calculated from the tracks in Fig. 2B. Solid lines were fitting results with . The fitted γR values are 0, 0.77, and 1.80 for gray, brown, and black tracks, respectively. D. Relationship between the rotational exponent γR and the mean y position of the bacterial trajectories. Errors denote SEM. Red and purple shaded area represent the region 3 μm from the LSW and RSW, respectively. Green shaded area represents the MA region. E&F. Normalized distribution of tumble angle for tracks along the sidewalls (E) and in the MA region (F). Red solid lines are fitting results with an exponential function: y = a ∗ exp (x/b). G. Distribution of bacteria along the y-axis in the lanes with a width of 44 μm. The shades of different colors denote the same meaning as in Fig. 2D. H. Drift velocity of bacterial cells in the three regions. I. Proportions of bacterial trajectories in the three regions. Errors are standard deviations (SD).

Revised version: A. Schematic drawing of three typical motion states. B. Trajectories for three typical motion states. Red dots represent the start of each trajectory. C. Calculating the rotational exponent for each trajectory, 𝛾𝑅, by fitting the mean-squared orientational displacement 〈MSOD〉. Dots were experimental data calculated from the tracks in Fig. 2B. Solid lines were fitting results with 〈MSOD〉 = 𝐶𝑅𝜏𝛾𝑅. The fitted 𝛾𝑅 values are 0, 0.77, and 1.80 for gray, brown, and black tracks, respectively. D. Relationship between the rotational exponent 𝛾𝑅 and the mean 𝑦 position of the bacterial trajectories. Errors denote SEM. Red and purple shaded area represent the region 3 μm from the LSW and RSW, respectively. Green shaded area represents the MA region. E&F. Normalized distribution of tumble angle for tracks along the sidewalls (E) and in the MA region (F). Red solid lines are fitting results with an exponential function: 𝑦 = 𝑎 ∗ exp (𝑥/𝑏). G. Distribution of bacteria along the y-axis in the lanes with a width of 44 μm. The shades of different colors denote the same meaning as in Fig. 2D. H. Drift velocity of bacterial cells in the three regions. I. Proportions of bacterial trajectories in the three regions. Errors are standard deviations (SD).

Earlier version: We distinguished these motion states by calculating the rotational exponent (γR) for each trajectory, fitting the mean-squared orientational displacement ⟨MSOD⟩ for τ < 0.4 s: where θ(t) is the angle between the swimming direction and the positive x-axis direction at time t, τ is the lag time, CR is a constant, and γR is the rotational exponent. Fig. 2C shows the γR values for the three trajectories in Fig. 2B, while Fig. 2D illustrates the relationship between γR and average y-coordinate. Trajectories further from sidewalls exhibit larger γR, indicating more curved paths, while γR decreases sharply as the cells approach sidewalls, representing straighter trajectories.

Revised version: We distinguished these motion states by calculating the rotational exponent (𝛾𝑅) for each trajectory, fitting the mean-squared orientational displacement 〈MSOD〉 for τ < 0.4 s: where 𝜃(𝑡) is the angle between the swimming direction and the positive x-axis direction at time t, 𝜏 is the lag time, 𝐶𝑅 is a constant, and 𝛾𝑅 is the rotational exponent. Fig. 2C shows the 𝛾𝑅 values for the three trajectories in Fig. 2B, while Fig. 2D illustrates the relationship between 𝛾𝑅 and average y-coordinate. Trajectories further from sidewalls exhibit larger 𝛾𝑅 , indicating more curved paths, while 𝛾𝑅 decreases sharply as the cells approach sidewalls, representing straighter trajectories.

Earlier version: We also analyzed tumble kinematics for cells swimming along and away from sidewalls. Fig. 2E and 2F show normalized tumble angle distributions for cells in sidewalls (SW) and middle areas (MA), respectively. Red solid lines represent exponential function fits y = a ∗ exp (x/b), with fitted b values of 0.25 rad (SW) and 1.35 rad (MA). Consistent with previous reports (41), we observed a higher occurrence of large reorientation angles for surface-swimming cells in the MA region (Fig. 2F). Interestingly, we also found large tumble angles in the SW region, albeit with a smaller characteristic angle than in the MA region. Movie S2 demonstrates an example of a large reorientation angle during tumbling.

Revised version: We also analyzed tumble kinematics for cells swimming along and away from sidewalls. Fig. 2E and 2F show normalized tumble angle distributions for cells in sidewalls (SW) and middle areas (MA), respectively. Red solid lines represent exponential function fits 𝑦 = 𝑎 ∗ exp (𝑥/𝑏), with fitted b values of 0.25 rad (SW) and 1.35 rad (MA). Consistent with previous reports (41), we observed a higher occurrence of large reorientation angles for surface-swimming cells in the MA region (Fig. 2F). Interestingly, we also found large tumble angles in the SW region, albeit with a smaller characteristic angle than in the MA region. Movie S2 demonstrates an example of a large reorientation angle during tumbling.

Earlier version: Given the strong spatial dependence of trajectory characteristics, we classified the cells into three distinct populations. The distribution of cells along the y-axis is shown in Fig. 2G. Based on mean distances from the sidewalls, we divided the lane into three areas (viewing along the positive x-axis): the left sidewall (LSW), the middle area (MA), and the right sidewall (RSW). As shown by different shaded colors in Fig. 2D and 2G, we defined cells in LSW, RSW and MA for trajectories with ⟨x⟩ ≤ d, ⟨x⟩ ≥ w − d, and d < ⟨x⟩ < w − d, respectively. The threshold d was set at 3 μm (approximately the average length of a cell body), and w represents the lane width. The proportion of trajectories in different regions is shown in Fig. 2I. For the gradient assay, the probabilities of trajectories being in LSW, MA and RSW regions were 0.20 ± 0.01, 0.59 ± 0.02, and 0.21 ± 0.03, respectively. In the control assay without gradient, these values were 0.18 ± 0.03, 0.70 ± 0.03, and 0.12 ± 0.01, respectively.

Revised version: Given the strong spatial dependence of trajectory characteristics, we classified the cells into three distinct populations. The distribution of cells along the y-axis is shown in Fig. 2G. Based on mean distances from the sidewalls, we divided the lane into three areas (viewing along the positive x-axis): the left sidewall (LSW), the middle area (MA), and the right sidewall (RSW). As shown by different shaded colors in Fig. 2D and 2G, we defined cells in LSW, RSW and MA for trajectories with 〈𝑥〉 ≤ 𝑑, 〈𝑥〉 ≥ 𝑤 − 𝑑, and 𝑑 < 〈𝑥〉 < 𝑤 − 𝑑, respectively. The threshold d was set at 3 μm (approximately the average length of a cell body), and 𝑤 represents the lane width. The proportion of trajectories in different regions is shown in Fig. 2I. For the gradient assay, the probabilities of trajectories being in LSW, MA and RSW regions were 0.20 ± 0.01, 0.59 ± 0.02, and 0.21 ± 0.03, respectively. In the control assay without gradient, these values were 0.18 ± 0.03, 0.70 ± 0.03, and 0.12 ± 0.01, respectively.

Earlier version: We then compared the drift velocity of cells in the three areas separately for both gradient and control assays. As shown in Fig. 2H, cells in the LSW region tended to display a negative drift with μm/s and -5.0 ± 2.0μm/s for gradient and control assays, respectively. Cells in the MA region exhibited minimal drift with μm/s and 0.4 ± 0.6 μm/s for gradient and control assays, respectively. Cells in the RSW region displayed a typical positive drift with μm/s and 6.2 ± 1.9 μm/s for gradient and control assays, respectively. Thus, cells in the RSW region dominated the difference in drift velocity between the two assays.

Revised version: We then compared the drift velocity of cells in the three areas separately for both gradient and control assays. As shown in Fig. 2H, cells in the LSW region tended to display a negative drift with vLSWd = -5.7±2.1 μm/s and -5.0 ± 2.0μm/s for gradient and control assays, respectively. Cells in the MA region exhibited minimal drift with vLSWd = -0.3 ± 1.2 μm/s and 0.4 ± 0.6 μm/s for gradient and control assays, respectively. Cells in the RSW region displayed a typical positive drift with vLSWd = 14.8 ± 1.2 μm/s and 6.2 ± 1.9 μm/s for gradient and control assays, respectively. Thus, cells in the RSW region dominated the difference in drift velocity between the two assays.

Earlier version: In summary, cells swim on the right-hand side under the effect of the bottom surface of lanes (9). The LSW and RSW cells swim down and up the gradient, respectively, while MA cells swim in circles without any drift. RSW cells demonstrated more persistent swimming towards the attractant in the L-aspartate gradient, exhibiting a higher positive drift velocity. Detailed information on the difference in drift velocity between LSW and RSW cells can be found in the supplemental materials. Notably, the attractant gradient minimally affects the ratio of bacteria in different regions.

Revised version: In summary, cells swim on the right-hand side under the effect of the bottom surface of lanes (9) , a behavior fundamentally dictated by the chirality of their surface interaction. The LSW and RSW cells swim down and up the gradient, respectively, while MA cells swim in circles without any drift. RSW cells demonstrated more persistent swimming towards the attractant in the L-aspartate gradient, exhibiting a higher positive drift velocity. To further illustrate the significance of this chirality, consider a hypothetical scenario where the bacteria exhibited counterclockwise circular swimming on the surface. In such a case, the hydrodynamic interaction would cause cells to veer to the left rather than the right. Consequently, the cells swimming up-gradient would accumulate along LSW, while those swimming down-gradient would accumulate along RSW. While the overall enhancement of chemotaxis via confinement would persist, the spatial distribution of the up- and down-gradient populations would be mirrored relative to the channel axis. Detailed information on the difference in drift velocity between LSW and RSW cells can be found in the supplemental materials. Notably, the attractant gradient minimally affects the ratio of bacteria in different regions.

The dependence of chemotactic performance on lane width

Having established that E. coli chemotaxis on lane surfaces is closely related to sidewall confinement, we investigated chemotaxis in lanes of different widths to mimic bacterial chemotactic behavior in interstitial tissues or tubes of varying sizes. We compared five lanes with widths of 6 μm, 8 μm, 10 μm, 25 μm, and 44 μm. Drift velocity was calculated using the same method as in Fig. 1C, with results shown in Fig. 3A. The mean drift velocities were 1.8 ± 1.3 μm/s, 7.5 ± 1.4 μm/s, 4.6 ± 1.7 μm/s, 2.5 ± 0.5 μm/s, and 2.6 ± 0.4 μm/s for lanes of width 6 μm, 8 μm, 10 μm, 25 μm, and 44 μm, respectively. Bacteria exhibited optimal chemotaxis (highest drift velocity) in lanes that were 8 μm wide. To rule out any potential ratchet effect caused by the asymmetric entrances, we performed control measurements using a setup with straight entrances; these results confirmed our findings and are shown in Fig. S3.

A. The drift velocity for the experiments in lanes of different widths. The error in drift velocity is the standard deviation. B. The distribution of the radius of circular swimming from 382 trajectories in the MA regions. The peak value of the radius is ∼10 μm. C. The swimming speed of cells in lanes with different widths. The mean swimming speeds and standard deviations for lanes with widths of 6 μm, 8 μm, 10 μm, 15 μm, 25 μm, and 44 μm are 22.7 ± 8.7 μm/s, 25.4 ± 7.7 μm/s, 27.1 ± 7.9 μm/s, 21.9 ± 7.8 μm/s, 22.1 ± 7.7 μm/s, and 22.6 ± 7.5 μm/s, respectively.

Noting that this optimal width is of the same order of magnitude as the radius of bacterial circular swimming on a solid surface, we measured the radius of circular swimming for cells in the MA region. As shown in Fig. 3B, the peak radius in the distribution of circular swimming radii is ∼10 μm, which is close to the lane width supporting optimal chemotaxis.

It has been reported that bacteria swimming in capillaries of different radii possess different swimming speeds (42). Therefore, we calculated the cell swimming speeds for all lane widths. As shown in Fig. 3C, swimming speed remains relatively constant across different lane widths. This indicates that the optimized chemotaxis is not a result of varying swimming speeds.

The optimal chemotaxis resulted from more cells moving towards the attractant along the right sidewall

To identify the underlying mechanism of optimal lane width for chemotaxis, we evaluated the drift velocity for cells in the LSW, MA, and RSW regions separately. The results are shown in Fig. 4A. Cells in the MA region did not exhibit effective drift, while LSW and RSW cells demonstrated negative and positive drift, respectively, consistent with the findings in Fig. 2H. The drift velocity for RSW cells increased with lane width, and remained relatively constant above a lane width of 8 μm, thus not exhibiting a highest value in the 8 μm-wide lanes. Therefore, the width-dependent optimal chemotaxis is not attributable to an optimal drift velocity for RSW cells.

A. Mean drift velocities in the LSW, MA, and RSW regions of lanes with different widths. The errors are SD. B. The proportion of bacterial cells in the LSW, MA, and RSW regions of lanes with different widths. Errors represent SD. C. The proportion of RSW-UG (cells swimming up-gradient in the RSW) and RSW-DG (cells swimming down-gradient in the RSW) for lanes with different widths.

We then sought a comparison of cell proportions in the three areas, as plotted in Fig. 4B. The proportion of cells in the MA region increased gradually with lane width due to the increase in area. Notably, the largest proportion of RSW cells appears in the 8 μm-wide lane. In other words, the lane width-dependent optimal chemotaxis primarily resulted from changes in cell numbers on the sidewalls. The relation between cell enrichment and lanes’ width was plotted in Fig. S4.

To further investigate changes in up- and down-gradient cell numbers in the RSW region, we calculated the proportion of RSW-UG (cells swimming up-gradient in the RSW) and RSW-DG (cells swimming down-gradient in the RSW) cells. As shown in Fig. 4C, more cells moved towards the attractant in the 8 μm-wide lane. Thus, the optimal chemotaxis in the 8 μm-wide lane resulted from a higher number of cells moving towards the attractant along the right sidewall.

Simulation of E. coli chemotaxis within lane confinements

To explore the relationship between the circular swimming of bacteria on lane surfaces and optimal chemotaxis, we simulated two-dimensional chemotaxis of bacteria with varying circular swimming radii on lanes of different widths in a steady linear gradient of L-aspartate.

Earlier version: In our simulation, cells were treated as self-propelled particles performing a random walk in a run-and-tumble mode within a two-dimensional space. Cells could be in either a run or tumble state, determined by the intracellular chemotaxis signal. During a run, cells swam smoothly at a constant speed of v0 =20 μm/s. The effect of the bottom surface on bacterial swimming was modeled by adding a constant angular velocity term to the swimming direction: where θ denotes the angle between the swimming direction and the positive x-axis, Δt is the time step, r is the radius of circular swim, Dr is the rotational diffusion coefficient, and n(0, 1) represents a random number drawn from a standard normal distribution. When cells collided with sidewalls, their velocities instantly aligned with the sidewall (43), and could potentially leave the sidewall after their next tumble. During a tumble, cells stopped moving, and their swimming direction changed. The change in θ was selected from the distribution of tumble angles observed in Fig. 2E and F for different areas.

Revised version: In our simulation, cells were treated as self-propelled particles performing a random walk in a run-and-tumble mode within a two-dimensional space. Cells could be in either a run or tumble state, determined by the intracellular chemotaxis signal. During a run, cells swam smoothly at a constant speed of 𝑣0 =20 μm/s. The effect of the bottom surface on bacterial swimming was modeled by adding a constant angular velocity term to the swimming direction: where θ denotes the angle between the swimming direction and the positive x-axis, Δ𝑡 is the time step, 𝑟 is the radius of circular swim, Dr is the rotational diffusion coefficient, and n(0, 1) represents a random number drawn from a standard normal distribution. When cells collided with sidewalls, their velocities instantly aligned with the sidewall (43), and could potentially leave the sidewall after their next tumble. During a tumble, cells stopped moving, and their swimming direction changed. The change in θ was selected from the distribution of tumble angles observed in Fig. 2E and F for different areas.

Initially, 100 cells were evenly distributed on the surface of a lane measuring 160 μm in length and w in width. A linear concentration gradient of L-aspartate, identical to the experimental conditions, was established along the x-axis according to: with the concentration c in μM and the position x in μm. The simulation for each set of parameters ran for 150 s with a timestep of 0.05 s and was repeated 50 times.

Earlier version: Figure 5A illustrates the relationship between chemotactic drift velocity and lane width for cells with a 10 μm-radius circular swim. The highest drift velocity was observed at a lane width of 8 μm. Figure 5B shows the proportions of cells in the three regions (LSW, MA, RSW) for different lane widths, presenting profiles similar to those observed experimentally in Fig. 4B. The drift velocity for 3 regions was shown in Fig. S5. Results for different swimming radii are presented in Fig. 5C, demonstrating that the optimal lane width for chemotaxis is closely related to the radius of circular swimming. It is important to note that our model treats bacteria as active particles and does not explicitly account for the steric exclusion of the flagella and cell body. In reality, these effects would restrict the cells’ ability to turn freely in narrow lanes (≤ 6 μm), thereby reducing the cell density mismatch between the LSW and RSW and decreasing the drift velocity compared to the simulation.

Revised version: Figure 5A illustrates the relationship between chemotactic drift velocity and lane width for cells with a 10 μm-radius circular swim. The highest drift velocity was observed at a lane width of 8 μm. Figure 5B shows the proportions of cells in the three regions (LSW, MA, RSW) for different lane widths, presenting profiles similar to those observed experimentally in Fig. 4B. The drift velocity for 3 regions was shown in Fig. S5. Results for different swimming radii are presented in Fig. 5C, demonstrating that the optimal lane width for chemotaxis is closely related to the radius of circular swimming. It is important to note that our model treats bacteria as active particles and does not explicitly account for the steric exclusion of the flagella and cell body. In reality, these effects would restrict the cells’ ability to turn freely in narrow lanes (≤ 6 μm), thereby reducing the cell density mismatch between the LSW and RSW and decreasing the drift velocity compared to the simulation. Consequently, the distinct behaviors of the left and right sidewall populations observed in the simulation at very small widths (Fig. 5B) reflect the idealized nature of our model, specifically the absence of steric hindrance.

Simulation of bacterial chemotaxis on the surface of lanes.

A. The relationship between drift velocity Vd of cells with 10 μm-radius circular swim and the lane width w. Errors denote SEM. B. The proportion of bacterial trajectories in the LSW, MA, and RSW regions of lanes with different widths, for cells with 10 μm-radius circular swim. C. The drift velocity Vd of cells with different radii of circular swim in lanes of different widths. D. The relationship between circular swim radius and lane width for optimal chemotaxis (maximal drift velocity) in Fig. 5C. The widths were extracted by fitting the peak value of w in Fig. 5A with a gaussian function. The red solid line represents a linear fit. The slope is 0.66 ± 0.03. Errors denote SD.

As shown in Fig. 5D, the relationship between optimal lane width (w) and circular swimming radius (r) was fitted with a linear function w = kr, with k = 0.66 ± 0.03. This indicates a clear proportional relationship between the lane width for optimal chemotaxis and the circular swimming radius of bacteria.

Geometrical analysis of optimal lane width for chemotaxis

We have established that the optimal lane width for chemotaxis (8 μm) primarily results from a higher proportion of cells in the RSW region swimming up-gradient. Additionally, this optimal width is directly proportional to the circular swimming radius of the bacteria. To further elucidate the underlying mechanism of this phenomenon, we employed geometrical analysis.

We assumed an initial uniform distribution of cells across the lane width. As shown in Fig. 6A-C, we divide the width of the lane into three scenarios:

Added in the revised version: 0 < 𝑤 ≤ 𝑟: For a cell at any point P in the channel, its swimming direction can be arbitrary. However, only cells swimming in directions between θ1 and θ2 (highlighted in green) can swim up-gradient in the RSW after their first collision with the sidewall. The proportion of these cells can be calculated as: where 𝑚 = 𝑤/𝑟 ≤ 1.

Added in the revised version: 𝑟 < 𝑤 ≤ 2𝑟 : Only the cells in the region y < 𝑟 and with velocity direction between θ1 and θ2 can swim up-gradient in the RSW after their first collision with the sidewall. The proportion of cells in this region is:

Added in the revised version: 𝑤 > 2𝑟: The proportion of cells that can swim up-gradient in the RSW is calculated the same as in case 2:

Geometrical analysis of optimal lane width for chemotaxis.

Earlier version: A. Case 1: 0 < w ≤ r. B. Case 2: r < w ≤ 2r. C. Case 3: w > 2r. The red arrows denote the velocity direction. The black solid lines represent the sidewalls. Dashed circles are trajectories. The up-gradient direction is along the positive x-axis. The green shade areas label the swim direction of cells that can swim up-gradient in the RSW. D. Relationship between probability (P) of cells swimming up-gradient in the RSW and m, where m = w/r. Red shaded area denotes m for maximal P, m ∈ (0.7,0.8).

Revised version: A. Case 1: 0 < 𝑤 ≤ 𝑟. B. Case 2: 𝑟 < 𝑤 ≤ 2𝑟. C. Case 3: 𝑤 > 2𝑟. The red arrows denote the velocity direction. The black solid lines represent the sidewalls. Dashed circles are trajectories. The up-gradient direction is along the positive x-axis. The green shade areas label the swim direction of cells that can swim up-gradient in the RSW. D. Relationship between probability (P) of cells swimming up-gradient in the RSW and m, where m = w/r. Red shaded area denotes m for maximal P, 𝑚 ∈ (0.7,0.8).

Removed in the revised version: 0 < w ≤ r: For a cell at any point P in the channel, its swimming direction can be arbitrary. However, only cells swimming in directions between θ1 and θ2 (highlighted in green) can swim up-gradient in the RSW after their first collision with the sidewall. The proportion of these cells can be calculated as: where m = w/r ≤ 1.

Removed in the revised version: r < w ≤ 2r : Only the cells in the region y < r and with velocity direction between θ1 and θ2 can swim up-gradient in the RSW after their first collision with the sidewall. The proportion of cells in this region is:

Removed in the revised version: w > 2r: The proportion of cells that can swim up-gradient in the RSW is calculated the same as in case 2:

Earlier version: The complete relationship between the proportion of cells that can swim up-gradient in the RSW (P) and m is plotted in Fig. 6D. The maximum P occurs when m ∈ (0.7,0.8), marked by the red shaded area. This value is consistent with the slope k in Fig. 5D and results in the optimal width of about 8 μm in Fig. 3A for a circular swimming radius of 10 μm.

Revised version: The complete relationship between the proportion of cells that can swim up-gradient in the RSW (P) and m is plotted in Fig. 6D. The maximum P occurs when 𝑚 ∈ (0.7,0.8), marked by the red shaded area. This value is consistent with the slope k in Fig. 5D and results in the optimal width of about 8 μm in Fig. 3A for a circular swimming radius of 10 μm.

Crucially, the observed asymmetry arises because the angular ranges that geometrically facilitate RSW-UG accumulation (Fig. 6A-C) coincide with the up-gradient direction. Because cells moving up-gradient experience suppressed tumbling, they can maintain the steady circular trajectories required to reach and align with the RSW. Conversely, while a pure geometric analysis suggests a symmetric potential for LSW-DG accumulation, these trajectories coincide with the down-gradient direction. Consequently, these cells experience enhanced tumbling, which distorts the circular trajectories. This prevents them from effectively reaching the LSW and also increases the probability of them leaving the wall.

Discussions

Bacterial behavior in confined environments differs significantly from that in free environments. Previous studies have reported that cells tend to swim in circles and accumulate near a solid-liquid interface, leading to diminished chemotaxis. However, recent research on chemotaxis of swarm cells suggests that crowded environments may enhance bacterial chemotaxis on surfaces. In this study, we examined the chemotaxis of E. coli on the bottom surface of microfluidic lanes, which simulate the spatial constraints imposed by neighboring cells in a swarm, exploring the effects of lane confinement on bacterial chemotaxis.

We found that in the presence of sidewalls, cells that would otherwise have been unable to exhibit chemotaxis showed significant chemotactic behavior in a linear gradient of L-aspartate. Cells in different regions of the lanes demonstrated distinct motion behaviors. In the MA region, cells showed no drift, similar to observations in chemotaxis near surfaces. Conversely, cells exhibited positive and negative drift when in the RSW and LSW regions, respectively.

Further investigation into chemotaxis in lanes of varying widths demonstrated that cells exhibited the highest chemotactic drift velocity in lanes that were 8 μm wide. This width is close to the radius of bacterial swimming circles. By analyzing the proportion and drift velocity of cells in different regions, we determined that this optimal lane width was primarily due to the higher proportion of cells swimming up-gradient along the right sidewall.

Earlier version: To explore the relationship between chemotactic drift velocity and circular swimming radius of bacteria, we conducted systematic simulations of E. coli chemotaxis with different circular swimming radii in lanes of varying widths under a linear concentration gradient of L-aspartate. These simulations explicitly incorporated the clockwise circular motion of surface-swimming bacteria and the alignment effect induced by sidewall collisions. We found that optimal chemotaxis occurred when the swimming radius was close to the lane width, with w = 0.66r. This suggests a significant proportional relationship between the optimal lane width for chemotaxis and the circular swimming radius of bacteria.

Revised version: To explore the relationship between chemotactic drift velocity and circular swimming radius of bacteria, we conducted systematic simulations of E. coli chemotaxis with different circular swimming radii in lanes of varying widths under a linear concentration gradient of L-aspartate. These simulations explicitly incorporated the chiral clockwise circular motion of surface-swimming bacteria and the alignment effect induced by sidewall collisions. We found that optimal chemotaxis occurred when the swimming radius was close to the lane width, with w = 0.66r. This suggests a significant proportional relationship between the optimal lane width for chemotaxis and the circular swimming radius of bacteria.

To elucidate the mechanism underlying this phenomenon, we employed geometric analysis. We demonstrated that bacteria in lanes could swim up-gradient in the RSW with maximal probability when the lane width was 0.7-0.8 times the radius of circular swimming. This finding is consistent with the experimental results presented in Fig. 4C.

Our findings have important implications for the screening of single-celled swimming organisms and the study of restricted chemotactic behavior. This research provides insights into how confinement affects bacterial chemotaxis and may inform the design of microfluidic devices for bacterial manipulation and analysis.

Bacteria frequently inhabit confined spaces such as soil pores, sediment interstices, rock crevices, biological tissues, and hydrogels. In these constrained environments, bacterial behavior often deviates significantly from that observed in unconfined conditions, typically resulting in reduced chemotactic ability. However, our findings challenge this notion, suggesting that certain types of environmental confinement may actually enhance bacterial chemotaxis. This discovery could contribute to our understanding of microbial ecology across various habitats, including soil, oceans, and the human body.

Furthermore, these insights could have potential applications. In biotechnology, optimizing confinement conditions could improve the efficiency of bioremediation processes or the production of microbial products. In medicine, understanding how pathogens navigate confined spaces within the body could lead to novel strategies for preventing or treating infections.

Materials and methods

Strain and plasmids

Removed in the revised version: The plasmid pTrc99a-mCherry expresses mCherry under an IPTG-inducible promoter. The wild-type E. coli K12 strain AW405 (HCB1) transformed with pTrc99a-mCherry was used for fluorescent tracking of cell trajectories. Cells were streaked on a Petri dish containing 1.5% agar and lysogeny broth (1% tryptone, 0.5% NaCl, 0.5% yeast extract, supplemented with the appropriate antibiotic: 100 μg/ml ampicillin). A single colony was inoculated into lysogeny broth and grown overnight at 33 °C with 200 rpm rotation.

Earlier version: The saturated culture was then diluted 1:100 (100 μl in 10 ml) into tryptone broth (1% tryptone, 0.5% NaCl, supplemented with the appropriate inducer and antibiotic: 0.1 mM IPTG, 100 μg/ml ampicillin) and grown at 33 °C with 200 rpm rotation to OD600=0.53. Cells were then harvested from culture media by centrifugation at 1.2×g for 6 min at room temperature. The pellet was resuspended by gently mixing in motility medium (10 mM potassium phosphate, 0.1 mM ethylenediaminetetraacetic acid (EDTA), 10 mM lactic acid, and 1 μM methionine at pH 7.0). The cells were washed three times to replace growth medium with motility medium. One hour after fluorescence maturation in a 33°C incubator, cells were placed in a 4°C refrigerator ready for microfluidic experiments.

Revised version: The plasmid pTrc99a-mCherry expresses mCherry under an IPTG-inducible promoter. The wild-type E. coli K12 strain AW405 (HCB1) transformed with pTrc99a-mCherry was used for fluorescent tracking of cell trajectories. Cells were streaked on a Petri dish containing 1.5% agar and lysogeny broth (1% tryptone, 0.5% NaCl, 0.5% yeast extract, supplemented with the appropriate antibiotic: 100 μg/ml ampicillin). A single colony was inoculated into lysogeny broth and grown overnight at 33 °C with 200 rpm rotation. The saturated culture was then diluted 1:100 (100 μl in 10 ml) into tryptone broth (1% tryptone, 0.5% NaCl, supplemented with the appropriate inducer and antibiotic: 0.1 mM IPTG, 100 μg/ml ampicillin) and grown at 33 °C with 200 rpm rotation to OD600=0.53. Cells were then harvested from culture media by centrifugation at 1.2×g for 6 min at room temperature. The pellet was resuspended by gently mixing in motility medium (10 mM potassium phosphate, 0.1 mM ethylenediaminetetraacetic acid (EDTA), 10 mM lactic acid, and 1 μM methionine at pH 7.0). The cells were washed three times to replace growth medium with motility medium. One hour after fluorescence maturation in a 33°C incubator, cells were placed in a 4°C refrigerator ready for microfluidic experiments.

Microfluidic device design and fabrication

A hydrogel-based microfluidic device containing lanes of different width was designed using L-edit based on previous work (42). The silicon mold with the positive relief features was fabricated using the standard soft lithography technique performed in our laboratory. Firstly, a 4-inch diameter silicon wafer was spin washed sequentially using water, acetone, methanol and isopropanol. The wafer was thoroughly dried on a 150 ℃ hot plate to ensure that organics has volatilized completely. Then, the wafer was treated with air plasma (Harrick plasma, PDC-002, 30W) for 10 min before the SU-8 2025 photoresist (Microchem) was spin-coated according to the products protocols. The spin-coated wafer was placed on a hot plate at 65℃ for 3 min followed by 6 min at 95℃, then exposed to UV light using a maskless lithography machine (Heidelberg, μMLA) at 800 mJ/cm2 with defocus was set to 0. The post-exposure bake was performed at 65℃ for 2 min followed by 6 min at 95℃. Development procedure lasted about 6 minutes with constant shaking in the SU-8 developer. The wafer was rinsed with fresh developer followed by isopropanol. Finally, an approximately 30-μm-high photoresist layer was created as the master mold. The mold was fixed on a glass petri dish using double-sided tape.

Polydimethylsiloxane (PDMS) and curing agent was mixed with a 10:1 ratio (Sylgard 184, Dow Corning) and centrifuged at 3500 rpm for 5 min to remove air bubbles. The mixture was poured over the mold in glass petri dish and then placed at 4℃ for an hour to avoid interference from very small air bubbles. The PDMS was cured at 80℃ for 3 h, then separated from the wafer and cut into pieces, and holes were punched for inlets and outlets. The punched PDMS was cleaned with adhesive tape. The cleaned PDMS was bonded to a glass slide after air plasma treatment (Harrick plasma, PDC-002, 30W). The fabricated device was baked at 80℃ overnight to establish covalent bonding.

The hydrogel agarose was used at 2% (w/v) concentration in motility medium to create diffusion-permeable walls. Agarose was slowly injected through the inlet port using a syringe pump at a constant flow rate of 2 µL/min, and the injection process was carried out in a drying oven at a constant temperature of 68 ℃.

Microscopy and data acquisition

The PDMS microfluidic chip was placed on a Nikon Ti-E inverted optical microscope equipped with a specific filter set, a 20× objective (Nikon, CFI S Plan Fluor ELWD, 20×, NA 0.45), and a CMOS camera (Hamamatsu, ORCA-Flash4.0, pixel size = 6.5 μm). We set the objective lens’s correction collar to 1.2 mm to match the glass slide thickness. A halogen lamp serves as the bright-field light source to help locate the observation area, while a xenon lamp serves as the excitation light source for fluorescence experiments. The depth of field is 3.7 μm.

Gradient profile calibration

Fluorescein solution (Sigma) was used for visualizing the chemical concentration field in the microfluidic device. 100 μM fluorescein and motility medium were flowed into the source and sink channels respectively, at a constant flow rate of 10 µL/min. One fluorescent image was captured per minute. To avoid the impact of photobleaching on the results, we only illuminated the sample during the imaging process.

Chemotaxis assays

The aspartic acid and motility medium solutions were respectively introduced into the source and sink channels, and left for more than 1 hour to allow aspartic acid to diffuse through the agarose wall and form a stable concentration gradient field. Then, bacteria were introduced into the cell channel at a rate of 5 μL/min using a syringe pump for a duration of 90 seconds. After the bacteria injection was complete, the microvalves were closed to seal the cell channel. Bacteria then performed chemotaxis due to the perception of the concentration gradient field. After waiting for 2 minutes, the shutter was opened to illuminate the fluorescent bacteria with the excitation light, and video recording began. During the video acquisition process, fluorescence excitation light was continuously illuminated onto the bacteria. For each recording, 9600 frames (2048×2048 pixels) were captured at a frame rate of 20 fps and an exposure of 50 ms. Although the long-term illumination caused a decrease in the brightness of the fluorescent bacteria, the position information of the bacteria could still be accurately distinguished within the video duration (8 min).

Data analysis

Tracking

Image analysis was performed with the ImageJ (National Institutes of Health) TrackMate plugin. We chose the Log detector with an estimated blob diameter of 70 μm and a threshold of 800 μm to effectively recognize bacteria, based on the diameter of bacteria and the use of a 20× objective. Then, we used the LAP Tracker and set the maximum distance for frame-to-frame linking to 50 μm.

Swimming speed

Removed in the revised version: The velocity is calculated using the fourth order central difference method by (44):

Earlier version: The magnitude of the velocity is given by . The speed of each trajectory is defined as the average of vi. The swimming speeds of bacteria in lanes of different widths are represented by the mean and standard deviation of all trajectories.

Revised version: The velocity is calculated using the fourth order central difference method by (44): The magnitude of the velocity is given by 𝑣𝑖 = |v⃗i|. The speed of each trajectory is defined as the average of 𝑣𝑖. The swimming speeds of bacteria in lanes of different widths are represented by the mean and standard deviation of all trajectories.

Calculation of the radius of circular motion on the solid surface

A total of 382 trajectories with a rotational exponent greater than 1.6 in the MA region of the w = 44 μm lanes were selected for analysis. The trajectories were smoothed using the Savitzky-Golay filter with a window size of 5. The least squares method was then applied to solve the linear equations, yielding the center and radius of the fitted circles. The principle of this process is briefly described as follows:

Earlier version: The standard equation of a circle can be written as: where (a, b) are the coordinates of the circle’s center, and r is the radius. Expanding and rearranging the standard equation of a circle, it can be rewritten as: where c = r2 − a2 − b2.

Revised version: The standard equation of a circle can be written as: where (𝑎, 𝑏) are the coordinates of the circle’s center, and 𝑟 is the radius. Expanding and rearranging the standard equation of a circle, it can be rewritten as: where c = 𝑟2 − 𝑎2 − 𝑏2.

Rewriting the above equation in matrix form, we have: where is the response vector, and is the design matrix. The coefficients are given by:

Tumble angle

Earlier version: A tumble is defined as points where the velocity is lower than the average velocity of the trajectory ⟨vi⟩ and shows a substantial drop from the previous time point. Specifically, a tumble is identified when: where n ≥ 2, α = 1/6 is the threshold, and t1 and tn represent the start and end times of a tumble, respectively.

Revised version: A tumble is defined as points where the velocity is lower than the average velocity of the trajectory 〈𝑣𝑖〉 and shows a substantial drop from the previous time point. Specifically, a tumble is identified when: where 𝑛 ≥ 2, α = 1/6 is the threshold, and 𝑡1 and 𝑡𝑛 represent the start and end times of a tumble, respectively.

Simulation of bacterial chemotaxis within lanes

Earlier version: In our simulation, we introduced the effects of bottom surface and sidewalls into the 2-dimensional chemotaxis of E. coli. The simulations were carried out in a rectangular area of size 160 μm × w, where w denotes the lane width. Cells were treated as self-propelled particles. They could sense and adapt to the ligand in the surrounding environment via the intracellular chemotaxis signaling pathway, which could be described by a coarse-grained model of bacterial chemotaxis (45, 46): where a, m, and c represent the activity of the receptor-kinase complex, the methylation level of receptors, and the ligand concentration, respectively. N denotes the number of chemoreceptors in a Monod-Wyman-Changeux cluster (47). The values of parameters we used were the same as before (4, 47, 48): N = 4.6, Koff = 1.7 μM, Kon = 12 μM, α = 1.7, m0 = 1.0, kR = 0.005 s−1, kB = 0.010 s−1.

Revised version: In our simulation, we introduced the effects of bottom surface and sidewalls into the 2-dimensional chemotaxis of E. coli. The simulations were carried out in a rectangular area of size 160 μm × w, where w denotes the lane width. Cells were treated as self-propelled particles. They could sense and adapt to the ligand in the surrounding environment via the intracellular chemotaxis signaling pathway, which could be described by a coarse-grained model of bacterial chemotaxis (45, 46): where 𝑎 , 𝑚 , and 𝑐 represent the activity of the receptor-kinase complex, the methylation level of receptors, and the ligand concentration, respectively. N denotes the number of chemoreceptors in a Monod-Wyman-Changeux cluster (47). The values of parameters we used were the same as before (4, 47, 48): 𝑁 = 4.6, 𝐾𝑜𝑓𝑓 = 1.7 μM, 𝐾𝑜𝑛 = 12 μM, 𝛼 = 1.7, 𝑚0 = 1.0, 𝑘𝑅 = 0.005 s−1, 𝑘𝐵 = 0.010 s−1.

Removed in the revised version: The concentration of CheY-P, which changes due to the receptor-kinase phosphorylation on CheY, could be calculated by Yp = 7.86a. Consequently, the CW bias (B) of flagellar motors would also vary (49):

Earlier version: It has been reported that the average run time in the near-surface region for wildtype cells is ∼2.0 s (corresponding to a run-to-tumble rate of ∼0.5 s-1) (11), and the steady-state motor CW bias is ∼0.15 without a gradient. The run-to-tumble rate k was set to be equivalent to the motor CCW-to-CW switching rate. As this rate increases linearly with motor bias B (50), it was calculated using . The switching rate from tumble to run was 5 s−1 due to the constant tumble duration of 0.2 s (51).

Revised version: The concentration of CheY-P, which changes due to the receptor-kinase phosphorylation on CheY, could be calculated by Yp = 7.86𝑎. Consequently, the CW bias (B) of flagellar motors would also vary (49): It has been reported that the average run time in the near-surface region for wildtype cells is ∼2.0 s (corresponding to a run-to-tumble rate of ∼0.5 s-1) (11), and the steady-state motor CW bias is ∼0.15 without a gradient. The run-to-tumble rate k was set to be equivalent to the motor CCW-to-CW switching rate. As this rate increases linearly with motor bias B (50), it was calculated using s−1. The switching rate from tumble to run was 5 s−1 due to the constant tumble duration of 0.2 s (51).

Earlier version: In this way, the motion states of each cell at any moment are determined. During a run, cells swim smoothly with a constant speed of 20 μm/s, influenced by rotational diffusion and the effect of the bottom surface on bacterial swimming: where θ is the angle between the swimming direction and the positive x-axis, v0 = 20 μm/s is the swimming speed, r is the radius of circular swimming due to the effect of the bottom surface, and n(0, 2DrΔt) represents a normal distribution with zero mean and variance 2DrΔt. The rotational diffusion coefficient was 0.062 rad2/s (52). During a tumble, cells stop swimming and a tumble angle is selected from the distribution that we measured (Fig. 2E). We also imposed a maximum detention time limit which, when combined with the variable tumble rate, results in an average wall residence time of approximately 2 s, matching our experimental observations (Fig. S6B). Cells will get a random y-coordinate at the other side when they swim through the x boundary.

Revised version: In this way, the motion states of each cell at any moment are determined. During a run, cells swim smoothly with a constant speed of 20 μm/s, influenced by rotational diffusion and the effect of the bottom surface on bacterial swimming: where 𝜃 is the angle between the swimming direction and the positive x-axis, 𝑣0 = 20 μm/s is the swimming speed, 𝑟 is the radius of circular swimming due to the effect of the bottom surface, and 𝑛(0, 2𝐷𝑟Δ𝑡) represents a normal distribution with zero mean and variance 2𝐷𝑟Δ𝑡. The rotational diffusion coefficient was 0.062 rad2/s (52). During a tumble, cells stop swimming and a tumble angle is selected from the distribution that we measured (Fig. 2E). We also imposed a maximum detention time limit which, when combined with the variable tumble rate, results in an average wall residence time of approximately 2 s, matching our experimental observations (Fig. S6B). Cells will get a random y-coordinate at the other side when they swim through the x boundary.